---
title       : Alpha and Omega
description : First and Last Chapter
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:MultipleChoiceExercise lang:python xp:50 skills:1 key:b454ac081a
## Asking difficult questions

How great is DataCamp on a scale of 1-10?

*** =instructions
- 3
- 10
- 6
- 200

*** =hint
It's the last one. 

*** =pre_exercise_code
```{python}
```

*** =sct
```{python}
msg_bad = "Try again!"
msg_success = "Correct. You're good at this!"
test_mc(4, [msg_bad, msg_bad, msg_bad, msg_success])
```